## Exercise 01.1

Degrees Fahrenheit ($T_F$) are converted to degrees Celsius ($T_c$) using the formula:

$$
T_c = 5(T_f - 32)/9
$$

Write a program to convert 78 degrees Fahrenheit to degrees Celsius and print the result.
Write your program such that you can easily change the temperature in Fahrenheit that you are converting from.

Use the variable name `T_c` for the temperature in degrees Celsius.

<!-- grade_id: cell-ca2d6af7fabf2918, points: 0, grade: false, solution: true -->
```python
T_f = 78

T_c=5*(T_f-32)/8
print(T_c)

```

Output:
```
28.75
```

<!-- grade_id: cell-a7542313c8cdf4ed, points: 0, grade: true, solution: false -->
```python
assert round(T_c - 25.55555555555555, 10) == 0.0
```

Output:
```
AssertionError                            Traceback (most recent call last)
~\AppData\Local\Temp\ipykernel_3588\3042612109.py in <module>
----> 1 assert round(T_c - 25.55555555555555, 10) == 0.0

AssertionError:
```

## Exercise 01.2

You have been tasked with developing a mortgage affordability test tool. A component is to compute the monthly 
interest payments for a range of scenarios.

Interest on a particular mortgage is charged at fixed rate above the Bank of England (BoE) 'official Bank Rate'. 
Interest is computed per annum, and interest payments spread equally over each month of the year.
Write a program that computes the interest payable each month, with variables for:

1. Loan principal (amount borrowed)
1. Official Bank Rate (percentage, expressed per annum)
1. Rate over the official Bank Rate (percentage, expressed per annum)

Test your program with a loan principal of £150,000, BoE rate of 0.5%, and rate over the 
BoE rate of 1.49%. Use the variable name `interest` for the monthly interest amount.

<!-- grade_id: cell-aaab53918bb9627d, points: 0, grade: false, solution: true -->
```python
# Assign data to variables
principal = 150000

# YOUR CODE HERE
BoE = 0.005
Bank = BoE + 0.0149
interest = (principal*Bank)/12
print(interest)
```

Output:
```
248.75
```

<!-- grade_id: cell-258e3a2c80f6a765, points: 0, grade: true, solution: false -->
```python
assert round(interest - 248.75, 10) == 0.0
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-9-1e4224306964> in <module>
----> 1 assert round(interest - 248.75, 10) == 0.0

AssertionError:
```

```python

```